<a href='https://ai.meng.duke.edu'> = <img align="left" style="padding-top:10px;" src=https://storage.googleapis.com/aipi_datasets/Duke-AIPI-Logo.png>

# Assignment 2

```python
# Run this before any other code cell
# This downloads the csv data files into the same directory where you have saved this notebook

import urllib.request
from pathlib import Path
import os
path = Path()

# Dictionary of file names and download links
files = {'Auto.csv':'https://storage.googleapis.com/aipi_datasets/Auto.csv',
        'copper-new.txt':'https://storage.googleapis.com/aipi_datasets/copper-new.txt'}

# Download each file
for key,value in files.items():
    filename = path/key
    url = value
    # If the file does not already exist in the directory, download it
    if not os.path.exists(filename):
        urllib.request.urlretrieve(url,filename)
```

## Question 1 [12 points]
For each of the following modeling problems, indicate which evaluation metric you would select to evaluate your model using the test dataset.  Please select ONE of the following choices for each problem: MSE, MAE, MAPE, Precision, Recall, Accuracy and explain in 1-2 sentences why you chose that metric.  

1)    I am creating a model to detect lung toxicity from chest x-rays.  Lung toxicity can be fatal if not detected and treated sufficiently early.    

2)    I am creating a model to predict power demand for the local electric utility.  I know that if my model is off slightly on a regular basis it is not a major issue, but if I am way off during a day of particularly low or high demand, it can cost the utlity a lot of money (e.g. because they have to start up expensive “peaker plants” to produce extra energy if I predict high demand, so I better not be wrong).  

3)    I am a day trader (meaning I buy and sell stocks on a daily basis) and I have built a stock price prediction model to support my trading.  I take a careful approach to my trading and I buy and sell frequent, small amounts of stock at a time – I do not place big bets.  Therefore, large errors of my model are not significantly worse than numerous small errors over time.  I need a metric which is easily understandable and fits my activity.  

4)    I am building a model to determine which students should be allowed to come on campus to classes vs. which should stay at home and watch via zoom.  My model aims to predict which students are likely to be healthy (positive class) vs. which are likely to be at risk of COVID (negative class).  

YOUR ANSWER HERE:

## Question 2 [20 points]
In the class lecture we stated that the mean squared error (MSE) of a model is the sum of the squared bias plus the variance (plus a component of irreducible error which we will ignore in this question for purposes of simplifying).  In this problem I would like you to derive this relationship via equations.  
 
As a starting point, we know the following:  

$$ MSE = E[(y-\hat{f}(x))^2] $$

$$ Bias = (y-E[\hat{f}(x)]) $$

$$ Variance = E[(\hat{f}(x)-E[\hat{f}(x)])^2] $$

Also, we know that by definition for a random variable Z:  

$$ Var[Z] = E[Z^2]-(E[Z])^2 $$


Use the above to derive the equation for MSE:  

$$ MSE = Bias^2 + Variance $$

A formal mathematical proof is not required but you must show the equations for each step of your derivation, and for each step you must include a brief comment describing what you are doing at that step.

YOUR ANSWER HERE:

Let's take note that the expected value of sums is equal to the sum of the expected values:

$$
E[x+y] = E[x]+E[y]
$$
Notably the $E[y]=f(x)$ for given value $y$:

Then

$$
\begin{multline}
MSE = E[(y-\hat{f}(x))^2] \\
= E[y^2] - 2\,E[y \times \hat{f}(x)] + E[\hat{f}(x)^2] \\
= f(x)^2 - 2\,f(x)\,E[\hat{f}(x)] + E[\hat{f}(x)^2] \\
= f(x)^2 - 2\,f(x)\,E[\hat{f}(x)] + E[\hat{f}(x)^2] + \{(E[\hat{f}(x)])^2)-(E[\hat{f}(x)])^2)\}\\
= \{f(x)^2 - 2\,f(x)\,E[\hat{f}(x)] + (E[\hat{f}(x)])^2\} + \{E[\hat{f}(x)^2]-(E[\hat{f}(x)])^2)\}\\
= \{f(x) - E[\hat{f}(x)]\}^2 + (E[\hat{f}(x)^2] - (E[\hat{f}(x)])^2) \\
= Bias^2 + Variance
\end{multline}
$$

## Question 3: Linear Regression Closed Form Solution [20 points]
In most cases where we are training a linear regression model we will use an iterative solving method such as gradient descent.  However, as discussed in class for small datasets it is feasible to calculate the closed form solution for the optimal coefficients.  In this question we will practice this approach to understand how it is done.  

Complete the below function `lr_closedform()` which first calculates the closed form solution for the coefficients (in this case, the bias term and a single coefficient) of a linear regression model.  The function should then apply your model on the data in `X_test` and return the predictions.

Your code for this question must be done in plain Python - **you may NOT use any packages** (you will notice none have been imported for you yet in this notebook).  You may also NOT use any AI assistants (e.g. GitHub Copilot) on this - I want you to think though yourself how to approach this.

Your function should work for any data of the same shape as the training data provided (3 observations, 1 feature + bias).  It does not need to work on inputs of any other size.

Hint: if it has been a while since you have thought about linear algebra (or have never done it), you may need to look up how to invert a matrix.  For this problem, due to the small size of the data, it is fairly straightforward.

```python
X_train = [[1,2],[1,3],[1,4]] # Training data inputs - 3 observations, each having a bias multiplier in first position (a '1') and a single feature
y_train = [5,6,8] # Training data targets
X_test = [[1,1],[1,4],[1,0.5],[1,6]] # Test input - 4 observations, each having a bias multiplier (a '1') in first position and a single feature
```

<!-- grade_id: cell-fabd942f1f25519e, points: 0, grade: false, solution: true -->
```python
def lr_closedform(X_train,y_train,X_test):
    '''
    Computes the closed form solution for the linear regression model relating the inputs X and outputs y

    Inputs:
        X_train(list): array of input training data stored as list of lists (including bias multiplier)
        y_train(list): target data
        X_test(list): array of input test data stored as list of lists (including bias multiplier)

    Returns:
        preds(list): predictions of your model on the test data
    '''

    ### BEGIN SOLUTION ###
    

    ### END SOLUTION ###
```

<!-- grade_id: cell-71f51679497aa0d2, points: 20, grade: true, solution: false -->
```python
# Test cell
preds = lr_closedform(X_train,y_train,X_test)
assert len(preds) == 4


assert round(preds[0],1) == 3.3
assert round(preds[1],1) == 7.8
assert round(preds[2],1) == 2.6
assert round(preds[3],1) == 10.8
```

## Question 4: Classification metrics and the ROC curve [10 pts]
### Question 4.0 [5 pts]
Run the cell below to display the actual target values, stored as `y`, for a classification modeling task, along with the probabilistic predictions from a classification model (the probability of the positive class) stored as `pred_probs`.  

```python
import numpy as np
import matplotlib.pyplot as plt

import warnings
warnings.filterwarnings("ignore")

# Actual y values saved as y:
np.random.seed(0)
y = np.round(np.random.rand(20)).astype(int)

# Predictions (probabilities of the positive class) saved as pred_probs:
np.random.seed(1)
pred_probs = np.round(np.random.rand(20),decimals=2)

print('Y values are {}'.format(y))
print('Predicted probabilities from the model are {}'.format(pred_probs))
```

Now, complete the below function `calc_metrics()`.  The function takes as inputs the true values `y` and the probabilistic predictions `pred_probs` and then calculates the following metrics on the predictions, using a threshold of >= 0.5 for predicting the positive class (1):
- Accuracy
- Recall
- Precision
- F1 Score

Your function should return the four metrics above for the positive class (1) as float values, in the order shown above.

You may use NumPy methods but **may NOT use any Scikit-Learn methods** to do this problem.


<!-- grade_id: cell-b8fd7cb2dde7c4e2, points: 0, grade: false, solution: true -->
```python
def calc_metrics(y,pred_probs):
    '''
    Calculates the accuracy, recall, precision and F1 score for a set of predictions and actuals using threshold >= 0.5 for positive class

    Inputs:
        y(np.array): array of actual target values
        pred_probs(np.array): array of probabilistic predictions of the positive class

    Returns:
        acc(float): accuracy of the predictions
        recall(float): recall of the predictions
        precision(float): precision of the predictions
        f1(float): F-1 score of the predictions
    '''
    
    ### BEGIN SOLUTION ###

   

    ### END SOLUTION ###
```

<!-- grade_id: cell-57ba507fba8c945f, points: 5, grade: true, solution: false -->
```python
# Test cell
acc,recall,precision,f1 = calc_metrics(y,pred_probs)
print('Your function calculated the metrics as follows:')
print('Accuracy: {:.3f}'.format(acc))
print('Recall: {:.3f}'.format(recall))
print('Precision: {:.3f}'.format(precision))
print('F1 Score: {:.3f}'.format(f1))

assert np.round(acc,2)==0.40
assert np.round(recall,2)==0.29
assert np.round(precision,2)==0.67
assert np.round(f1,2)==0.40
```

### Question 4.1 [5 pts]
Now, complete the function `calc_roc()` to calculate the values needed to plot a ROC curve for the model.  The function should take as input the actual values `y` and the probabilistic predictions of the positive class `pred_probs`.  The function should vary the threshold for predicting the positive class (1) using the range `thresh_vals` provided as input and for each threshold value, should compute the true positive rate (tpr) and false positive rate (fpr) of the predictions.  Your function should return two lists: the first list containing the tpr values at each threshold, and the second list containing the fpr values at each threshold.


Again you **may NOT use any Scikit-Learn methods** for this part


<!-- grade_id: cell-1c75a807a7564e9f, points: 0, grade: false, solution: true -->
```python
def calc_roc(y,pred_probs,thresh_vals):
    '''
    Calculates the TPR and FPR of the predictions for all threshold values specified in the input.  Use >= threshold for 
    predicting the positive class

    Inputs:
        y(np.array): array of actual target values
        pred_probs(np.array): array of probabilistic predictions of the positive class
        thresh_vals(np.array): array of threshold values

    Returns:
        tprs(list): list of TPR values for each threshold
        fprs(list): list of FPR values for each threshold
    '''

    ### BEGIN SOLUTION ###

   

    ### END SOLUTION ###
```

<!-- grade_id: cell-3271899b8a665e16, points: 5, grade: true, solution: false -->
```python
# Test cell
thresh_vals = np.arange(0.0,1.01,0.01)
tprs,fprs = calc_roc(y,pred_probs,thresh_vals)
assert (len(tprs),len(fprs))==(101,101)
plt.plot(fprs,tprs,label='Model')
plt.plot([0,1],[0,1],linestyle='--',label='Random guessing')
plt.xlabel('fpr')
plt.ylabel('tpr')
plt.legend()
plt.show()

assert (np.round(np.mean(tprs),2)==0.39 or np.round(np.mean(tprs),2)==0.38)
assert (np.round(np.mean(fprs),2)==0.34 or np.round(np.mean(fprs),2)==0.33)
```

## Question 5: Bias-Variance Tradeoff [18 points]
In this question we are going to be working with an automotive dataset.  Our goal is to build a model to predict fuel efficiency (mpg) from engine horsepower.  We will be using a Random Forest for our model (you don't need to know what this is for purposes of answering this question, we will cover it later).    

```python
import pandas as pd
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_absolute_error
from sklearn.ensemble import RandomForestRegressor
from sklearn.model_selection import KFold

import warnings
warnings.filterwarnings("ignore")

# Read in data
data = pd.read_csv('Auto.csv')
# Remove rows with missing values
data = data[data['horsepower'] != '?']

# Define the features and target (X and y)
X = data['horsepower'].astype(int).values.reshape(-1,1)
y = data['mpg'].values.astype(float)

# Split into training and test sets
X_train_full,X_test,y_train_full,y_test = train_test_split(X, y, random_state=0,test_size=0.1)

# Model to use. Will vary min_samples_leaf and calculate cross-validation MAE 
model = RandomForestRegressor(min_samples_leaf=1,random_state=0)

# Values of the hyperparameter min_samples_leaf to compare using cross-validation
msl_range = np.arange(1,41,1)

# Number of folds to use for k-folds cross validation
folds = 3
```

### Q5.0 (9 points)

**When the value of min_samples_leaf is large (e.g. 35-40), is the model underfitting, overfitting, or fitting just right?  Defend your answer with 1-2 sentences and include a visualization of the model training and validation MAE generated by your code to explain your answer.**


Our task is to compare the performance of the model using the different values for the hyperparameter `min_samples_leaf` listed in `msl_range`.  You should compare them using **3-fold cross-validation** and you should use **Mean Absolute Error** as your evaluation metric for comparision. 

Use the below space to write your code to compare the performance of the model using the different values of `min_samples_leaf` and generate any charts needed to answer the question.  

YOUR ANSWER HERE:

<!-- grade_id: cell-aa448a45754b881a, points: 0, grade: false, solution: true -->
```python
# Use this cell to write any code and generate the visualizations you need to answer the question

### BEGIN SOLUTION ### 


### END SOLUTION ###
```

### Q5.1 (9 points)
When the value of min_samples_leaf is small (e.g. 1-5), is the model underfitting, overfitting, or fitting just right?  Defend your answer with 1-2 sentences and include a visualization of the model training and validation MAE generated by your code to explain your answer.

YOUR ANSWER HERE:

```python
# Use this cell to write any code and generate the visualizations you need to answer the question
```

## Question 6 [20 points]
We are given a dataset from an experiment looking at how the thermal expansion coefficient of copper changes with temperature.  Our task is to model the relationship of thermal expansion coefficent to temperature to be able to predict the coefficient for any new temperature value.  We will use a Linear Regression model, but will apply polynomial transforms to the input data X in order to better capture the relationship between X and y.  We will evaluate different polynomial transforms to select the optimal polynomial order to use.

```python
from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import PolynomialFeatures

# Load the data
copperdata = pd.read_csv('copper-new.txt',header=None)
copperdata['X'] = copperdata.apply(lambda x: x.str.split()[0][1],axis=1)
copperdata['y'] = copperdata.apply(lambda x: x.str.split()[0][0],axis=1)
copperdata = copperdata[['X','y']].astype(float)
copperdata.head()

```

### Question 6.0 [10 pts]
If we were to fit a linear regression model on this data, does it violate the assumption of linearity between the input and targets?  For credit you must justify your answer by referring to BOTH a scatterplot of the data and a plot of the residuals/errors vs your predictions and explain how we answer this question based on the plots. 


YOUR ANSWER HERE:

<!-- grade_id: cell-7f30b18bc352cccb, points: 0, grade: false, solution: true -->
```python
# Use this cell to write the code and generate the visualizations you need to answer the question

### BEGIN SOLUTION ###




### END SOLUTION ###
```

### Question 6.1 [10 pts]
Let's see if we can use polynomial regression to find a better fitting model than a standard linear regression.  When doing your polynomial transformation be sure to include a bias term

Complete the below function `polynomial_model()` which finds the optimal degree polynomial transform (from the list provided as input) using a fixed validation set.  Split your data to use 10% as the test set.  Of the remaining 90%, use 80% of that for training and 20% for validation.  Be sure to use `random_state=0` when splitting your data.  Your function should return the optimal degree of polynomial transform which minimizes MAE on the validation set, and the MAE of your final model on the test set (be sure to retrain on the full data after finding your optimal transform).


<!-- grade_id: cell-4b8c42820317ca0d, points: 0, grade: false, solution: true -->
```python
def polynomial_model(df,n_range):
    '''
    Finds the optimal degree of polynomial transform and calculates the MAE of the optimal model on the test set

    Inputs:
        df(DataFrame): dataframe containing the inputs and targets
        n_range(list): list of values to use for degree of the polynomial transform

    Returns:
        opt_degree(int): optimal degree of polynomial transform that minimizes MAE on the validation set
        test_mae(float): MAE on the test set
    '''
    
    ### BEGIN SOLUTION ###
    

    ### END SOLUTION ###
```

<!-- grade_id: cell-2a1cd5ba21767636, points: 10, grade: true, solution: false -->
```python
# Test cell
n_range = range(1,8)
opt_degree, test_mae = polynomial_model(copperdata, n_range)


print('Your function found the optimal degree was {}'.format(opt_degree))
print('Your function calculated a test set MAE of {:.3f}'.format(test_mae))

assert opt_degree == 5
assert np.round(test_mae,2)==0.16

```